<a href="https://colab.research.google.com/github/venkateshpaturu/venkatesh-HDS/blob/main/week02/week02_assignment.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 2 Programming Exercises

See: _McKinney 2.3_ and [Python Documentation](https://docs.python.org/3/tutorial/introduction.html) section 3.1.1 on using Python as a calculator.

**At the begining of the semester, all of the workshop programming exercises will be structured a specific way to make it easier to verify for yourself that you're on the right track as well as easier for me to do a first pass on automated grading.  The structure of each question will require you to write a function using Python code. Don't worry that we haven't talked about functions yet. Just edit the code between** `### BEGIN SOLUTION` and `### END SOLUTION` **as shown in the example below.**


**WHAT I PROVIDE:**
```
def some_function(parameter1, parameter2):

   ### BEGIN SOLUTION
   x = -1
   ### END SOLUTION
   
   return x
```

**WHAT YOU SHOULD DO:** Just change the parameter names (if you feel you need to) and the calculations between `### BEGIN SOLUTION` and `### END SOLUTION`.  This is just a made up example.
```
def some_function(a, b):

   ### BEGIN SOLUTION
   temp = a + b
   x = temp / a * b
   ### END SOLUTION
   
   return x
```


---
---

**Below each programming exercise are some tests (`assertions`) that verify your code is working correctly.  If any assertions fail, you know that something isnt' right with your code, but having all assertions pass doesn't necessarily mean your code is perfect, yet. You should also create your own tests to make sure your code is correct.**

**For now, please don't change any function names**

### 5.1 BMI Calculation

Create new names for the variables that you need in the code framework below and use the variables to compute the corresponding body mass index (BMI). You can use any meaningful variable names you like. When you compute the BMI, assign a variable to it, and then print the result to the notebook's output.

* Weght in kilograms is 90.7
* Height in meters in 1.82

Note that a simple BMI calculation is weight in kilograms divided by the square of height in meters.

$ bmi = \frac{weight\,\mathrm{kg}}{(height\,\mathrm{m})^2} $

For this problem, you'll be graded both on proper calculation as well as on programming style and appropriate variables names.

_Also note that BMI should not be considered a universal measure of health or fitness. A health BMI can vary greatly depending on many other fitness conditions and body types._

<!-- grade_id: bmi-answer, points: 0, grade: false, solution: true -->
```python
def bmi(weight_kg, height_m):

    # 1. Change the names of parameter1 and parameter2 to be something more meaningful
    # 2. Change the name of the return variable below

    ### BEGIN SOLUTION
    answer = -1
    ### END SOLUTION

    return answer
```

Test your code below.  The BMI for a weight of 90.7 kg and height of 182 cm should be approximately 27.38

<!-- grade_id: bmi-test, points: 0, grade: true, solution: false -->
```python
# Define the BMI function
def bmi(weight_in_kilograms, height_in_meters):
    body_mass_index = weight_in_kilograms / (height_in_meters ** 2)
    return body_mass_index

# print result
result = bmi(90.7, 1.82)
print(result)

assert round(bmi(90.7, 1.82), 2) == 27.38
assert round(bmi(91, 1.80), 2) == 28.09
assert round(bmi(70, 1.70), 2) == 24.22
```

Output:
```
27.38195870064002
```

```python
bmi(90.7, 1.82)
```

### 5.2 Weight from BMI

For the next operation, we're going to calculate Weight in kilograms, given someone's BMI and their height in meters.

This calculation will be:  $ weight\,\mathrm{kg} = {bmi}\times{(height\,\mathrm{m})^2} $

```python
def weight(bmi, height_m):

    ### BEGIN SOLUTION
    w = -1
    ### END SOLUTION

    return w
```

```python
# Define the function to calculate weight in kilograms based on BMI and height in meters
def weight_from_bmi(bmi, height_in_meters):
    weight_in_kilograms = bmi * (height_in_meters ** 2)
    return weight_in_kilograms

# weight result
weight_result = weight_from_bmi(27.38, 1.82)
print(weight_result)

# functions assertions test
assert round(weight_from_bmi(27.38, 1.82), 2) == 90.69
assert round(weight_from_bmi(28.09, 1.80), 2) == 91.01
assert round(weight_from_bmi(24.22, 1.70), 2) == 70.00
```

Output:
```
90.693512
```

```python

```

### 5.3 Height from BMI


For the next operation, we're going to calculate height in meters, given someone's BMI and their weight in kilograms.

This calculation will be:  $ {height\,\mathrm{m}} = \sqrt{\frac{weight\,\mathrm{kg}}{bmi}} $

For this equation, you'll notice that we need to calculate a square root.  Python has a function for this of course, `sqrt`, but that function is not part of the base Python language. Some Python features have to be imported.  The first line of code below imports the `sqrt` function from the [math](https://docs.python.org/3/library/math.html) library.


```python
from math import sqrt
def height(bmi, weight_kg):

    ### BEGIN SOLUTION
    h = -1
    ### END SOLUTION

    return h
```

```python
# Import the sqrt function from the math library
from math import sqrt

# Defining the function to calculate height and weight
def height_from_bmi(weight_in_kilograms, bmi):
    height_in_meters = sqrt(weight_in_kilograms / bmi)
    return height_in_meters

#  calculation
height_result = height_from_bmi(90.7, 27.38)
print(height_result)
```

Output:
```
1.8200650980900683
```

```python

```

### 5.4 BMI using Feet, Inches, and Pounds

Now, let's try to to create a function that knows how to first convert between imperial units (US inches and pounds) and SI units before calculating the bmi.  For these calculations use:

* $ 1\,\mathrm{inch} = 2.54\,\mathrm{cm} $
* $ 1\,\mathrm{lb} = 0.45454545\,\mathrm{kg} $


```python
def bmi_us(weight_lb, height_ft, height_in):

    ### BEGIN SOLUTION
    answer = -1
    ### END SOLUTION

    return answer
```

```python
# Define the conversion constants
INCH_TO_CM = 2.54
LB_TO_KG = 0.45454545

# Defining the function to calculate BMI using  units (inches and pounds)
def bmi_imperial_to_si(weight_in_pounds, height_in_inches):
    # Convert weight from pounds to kilograms
    weight_in_kilograms = weight_in_pounds * LB_TO_KG

    # Convert height from inches to meters
    height_in_meters = (height_in_inches * INCH_TO_CM) / 100

    # Calculate BMI using the SI formula
    body_mass_index = weight_in_kilograms / (height_in_meters ** 2)
    return body_mass_index

# calculation
bmi_result_imperial = bmi_imperial_to_si(200, 70)
print(bmi_result_imperial)  # This will print the BMI result
```

Output:
```
28.757014554845437
```

### 5.5 Target BMI
Assuming that someone doesn't change height over time, what is the approximate number of days it will take for someone to get from their height and weight to a BMI of 24.9 (typically regarded as being in a healthy range) while losing weight at a rate of 0.75 kg per week.

As needed, use variables and multiple calculation steps to make your thought process visible in the code. For instance, rather than using the notebook as a calculator, use variables for intermediate answers.

Again, you will receive points both for correct answers as well as for style.

In addition to the bmi, height, and weight equations we already defined above, here is some of the algebra to help you along:

$ height\,\mathrm{m} = \frac{height\,\mathrm{cm}}{100} $

$ days = {weeks} \times {7} $

$ weeks\,required\,to\,change\,weight = \frac{current\,weight - target\,weight}{weight\,change\,per\,week} $

<!-- grade_id: weight-answer, points: 0, grade: false, solution: true -->
```python
def lose_weight(weight, height, kg_per_week, target_bmi):

    # Note that the variables you need area already named for you:
    # height is the height in cm
    # weight is the weight in kg
    # kg_per_week is how much the person wants to lose per week
    # target_bmi is the intended target bmi

    ### BEGIN SOLUTION
    days = -1
    ### END SOLUTION

    return days
```

Test your code below.  It should return approximately 76.73 days for the given parameters above.

<!-- grade_id: weight-tests, points: 5, grade: true, solution: false -->
```python
# Define the function to calculate the target weight for a BMI of 24.9
def target_weight_for_bmi(height_in_meters, target_bmi=24.9):
    target_weight = target_bmi * (height_in_meters ** 2)
    return target_weight

# Defining the function to calculate the number of weeks to reach target weight
def weeks_to_reach_target(current_weight, target_weight, weight_loss_per_week):
    weight_difference = current_weight - target_weight
    weeks_required = weight_difference / weight_loss_per_week
    return weeks_required

# Defining the function to calculate the number of days to reach target BMI
def days_to_reach_bmi(current_weight_kg, height_cm, target_bmi=24.9, weight_loss_per_week=0.75):
    # Convert height from cm to meters
    height_m = height_cm / 100

    # Calculate the target weight for the target BMI
    target_weight = target_weight_for_bmi(height_m, target_bmi)

    # Calculate the weeks required to reach the target weight
    weeks_required = weeks_to_reach_target(current_weight_kg, target_weight, weight_loss_per_week)

    # Calculate the number of days (1 week = 7 days)
    days_required = weeks_required * 7
    return days_required

# calculation
current_weight_kg = 90.7
height_cm = 182

days_result = days_to_reach_bmi(current_weight_kg, height_cm)
print(days_result)
```

Output:
```
76.73157333333342
```

```python
lose_weight(90.7, 182, 0.75, 24.9)
```

---

### 5.6 Your choice

Every one of us has done projects where some kind of simple mathematical calculators are required.  Think about something recent from your field of expertise or an area of interest. Choose something that you find interesting and write the Python code required to do that calculation.  These examples show the level of complexity you need to have (not very complex).

* $ bmr_{male} = 66 + (13.7 \times weight\,\mathrm{kg}) + (5 \times height\,\mathrm{m}) - (6.8 \times age\,\mathrm{yr}) $

* $ bmr_{female} = 655 + (9.6 \times weight\,\mathrm{kg}) + (1.8 \times height\,\mathrm{cm}) - (4.7 \times age\,\mathrm{yr}) $

* $ Batting\,Average = \frac{Total\,Hits}{Total\,at\,Bats} $

* $ Area\,of\,circle = \pi r^2 $

* $ F = \frac{1}{2} m v^2 $

Then provide the code for how you would use variables and example values to do the calculation.  You do not to wrap this in a Python function. Just include the code and variables as I've shown below with the BMI calculation that we've been using. As the last command in the cell, print out the value of your calculation as shown in my example.


```python
def my_function(a, b):
    c = a ** b
    return c
```

```python
my_function(1, 8)
```

Output:
```
1
```

```python
my_function(2, 8)
```

Output:
```
256
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.




** NEW INSTRUCTIONS For Submitting your work**
First run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Use the Colab feature to save to your Github:
  1. You will want to go to the File Menu.
  2. Select Save a Copy in Github.
  3. Then Change the Repository to your personal repository.
  4. Click OK

You can verify that everything worked by going to your Github and viewing the file in place.





## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

2. Then open a new command line by clicking the `+` icon above the file list and chosing `Terminal`
3. At the command line in the new Terminal, follow these steps:
  1. Change directories to your project folder and the week02 subfolder (`cd <folder name>`)
  2. Make sure your project folders are up to date with github.com (`git pull`)
  3. Add the homework files for this week (`git add <file name>`)
  4. Commit your changes (`git commit -m "message"`)
  5. Push your changes (`git push`)
  
**Here's a full examle**
```
cd hds5210-2022/week02
git pull
git add week02_assignment_1.ipynb week02_assignment_2.ipynb
git commit -a -m "Submitting homework assignments for week 2"
git push
```

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.


---

If the message above says something like _Submitting the week 2 programming assignment_ or _Everything is up to date_, then your work was submitted correctly.